---
id: 64de1469fccce61940f498c4
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

The `rods` dictionary will represent the three rods with their disks. Give it the strings `'A'`, `'B'`, and `'C'` as keys and set each of them to an empty list.

# --hints--

Словник `rods` повинен мати ключ `'A'`.

```js
({ test: () => assert(__pyodide.runPython(`'A' in __locals.get("rods")`)) })

```

`rods['A']` має бути порожнім списком.

```js
({ test: () => assert(__pyodide.runPython(`
    a = __locals.get("rods")
    a['A'] == []
  `))
})
```

Словник `rods` повинен мати ключ `'B'`.

```js
({ test: () => assert(__pyodide.runPython(`'B' in __locals.get("rods")`)) })

```

`rods['B']` має бути порожнім списком.

```js
({ test: () => assert(__pyodide.runPython(`
    a = __locals.get("rods")
    a['B'] == []
  `))
})
```

Словник `rods` повинен мати ключ `'C'`.

```js
({ test: () => assert(__pyodide.runPython(`'C' in __locals.get("rods")`)) })

```

`rods['C']` має бути порожнім списком.

```js
({ test: () => assert(__pyodide.runPython(`
    a = __locals.get("rods")
    a['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
