---
id: 6557709b0aee699a6a00528c
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Now, replace the existent keys with the strings `A` and `B` — one for each node. Then, replace each value with the string representing the node connected to the key.

# --hints--

Словник повинен мати ключ `A`.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    "A" in graph
  `))
})
```

Ключ `A` повинен мати значення `B`.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph["A"] == "B"
  `))
})
```

Словник повинен мати ключ `B`.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    "B" in graph
  `))
})
```

Ключ `B` повинен мати значення `A`.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph["B"] == "A"
  `))
})
```

Словник повинен мати два ключі.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    len(graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
