---
id: 65789506b30453080f77470c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

So far, you have already met different data types:

- Immutable data types, such as integers, strings, tuples, and Booleans.
- Mutable data types, such as lists.

A dictionary is a mutable data type and it is identified by a pair of curly braces, `{}`.

Start by creating a variable called `copper` and assign it an empty dictionary using a pair of curly braces, in the same way you would create an empty list with a pair of square brackets.

# --hints--

Ви повинні мати змінну під назвою `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Змінна `copper` повинна мати значення порожнього словника. Use a pair of curly braces for that.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
